<a href="https://colab.research.google.com/github/RafaelNovais/MasterAI/blob/master/Assignment_2_NLP.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
from google.colab import drive
drive.mount('/content/drive')
```

Output:
```
Mounted at /content/drive
```

```python
NAME = "Rafael Novais de Melo = 23113607"
COLLABORATORS = ""
```

---

<img src="logo.png" width="250" height="250">

# CT5120/CT5146 Assignment 2

This task is focused on sentiment classification and span extraction from tweets. Please complete this task and upload your answers to Canvas as an iPython Notebook ONLY.
This assignment is due by 23:59 on November 18th 2024. Late submissions will be penalised by 1% for each day after this date. This is an individual assignment and your work must be your own.

You may use libraries such as SciKit-Learn to complete this assignment.

## Summary

In this task, given a tweet (text) we have two objectives:

1. Sentiment classification: Classify the tweet into one of three classes (positive, negative, neutral).

2. Sentiment span extraction: Extract the sequence of words from the tweet that expresses the given sentiment.

Consider the following example from the train dataset:

| textID | text | selected_text | sentiment |
| ------ | ---- | ------------- | --------- |
| 266b8792a0 |Just broke my favorite necklace  superglue? | Just broke my favorite necklace | negative |
| 8f3e73cf09 | "Screw the reviews, I thought Wolverine was awesome. But not enough Dominic Monaghan for my liking." | I thought Wolverine was awesome. | positive |
|... | ... | ... | ... |
| 266b8792a0 |Just broke my favorite necklace  superglue? | Just broke my favorite necklace | negative |


The dataset is divided into `train`, `dev` and `test` sets. The `train` set is used for the model training, while the `dev` set is used for validation and hyperparameter tuning.

```python
!pip install pandas
!pip install matplotlib
!pip install -U scikit
!pip install gensim
```

Output:
```
Requirement already satisfied: pandas in /usr/local/lib/python3.10/dist-packages (2.2.2)
Requirement already satisfied: numpy>=1.22.4 in /usr/local/lib/python3.10/dist-packages (from pandas) (1.26.4)
Requirement already satisfied: python-dateutil>=2.8.2 in /usr/local/lib/python3.10/dist-packages (from pandas) (2.8.2)
Requirement already satisfied: pytz>=2020.1 in /usr/local/lib/python3.10/dist-packages (from pandas) (2024.2)
Requirement already satisfied: tzdata>=2022.7 in /usr/local/lib/python3.10/dist-packages (from pandas) (2024.2)
Requirement already satisfied: six>=1.5 in /usr/local/lib/python3.10/dist-packages (from python-dateutil>=2.8.2->pandas) (1.16.0)
Requirement already satisfied: matplotlib in /usr/local/lib/python3.10/dist-packages (3.7.1)
Requirement already satisfied: contourpy>=1.0.1 in /usr/local/lib/python3.10/dist-packages (from matplotlib) (1.3.0)
Requirement already satisfied: cycler>=0.10 in /usr/local/lib/python3.10/dist-packages (from matplotlib) (0.12.1)
Requi
```

```python
import pandas as pd
from matplotlib import pyplot as plt
import numpy as np
import inspect
import re
import gensim.downloader
from sklearn.ensemble import AdaBoostClassifier
```

```python
train_df = pd.read_csv('/content/drive/MyDrive/Colab Notebooks/NLP/train.csv')
dev_df = pd.read_csv('/content/drive/MyDrive/Colab Notebooks/NLP/dev.csv')
test_df  = pd.read_csv('/content/drive/MyDrive/Colab Notebooks/NLP/test.csv')
train_df.columns
```

---
# 1. Data Analysis

## Task 1a
`(5 Marks)`



<!-- grade_id: class_frequency, points: 0, grade: false, solution: true -->
```python

def class_frequency(train_df, dev_df):
    """
    This function should takes train and dev dataframe and returns a dictionary where:

    Keys = sentiment labels i.e. {negative, positive, neutral}
    Values = (count of Key in train_df, count of Key in dev_df)
    Arguments:
    train_df: DataFrame for training dataset
    dev_df: DataFrame for dev dataset

    Returns
    (order does not matter)
    {
    'negative': (no_in_train, no_in_dev),
    'positive': (no_in_train, no_in_dev),
    'neutral':  (no_in_train, no_in_dev)
    }


    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: correct_output_type, points: 2, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
class_freq_outputs = class_frequency(train_df, dev_df)
assert len(class_freq_outputs) == 3
assert set(class_freq_outputs) == set(['negative', 'positive', 'neutral'])


```

<!-- grade_id: correct_output_value, points: 3, grade: true, solution: false -->
```python

```

# TASK 1B
Plot a group bar plot to show the distribution of sentiment classes (positive, negative, neutral) in the train and dev dataset.

As shown in the following illustration. `(10 marks)`
    
    
   <img src="https://github.com/gauneg/gauneg.github.io/blob/main/ds1.jpg?raw=true" alt="PLOT EXAMPLE" width="400px"/>

<!-- grade_id: bar_graph_viz, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

---
# 2. Sentiment Classification

## Task 2a

Train a sequence classification model. The model should take the `text` as input and return one out of 3 sentiment classes: `negative`, `neutral` or `positive`.
<img src="https://github.com/gauneg/gauneg.github.io/blob/main/sentiment_classification.png?raw=true" alt="PLOT EXAMPLE" width="500px"/>

INPUT : (`text`)

OUTPUT: `sentiment_class` i.e. one class from `(positive, negative, neutral)`


* You should have the data in the training and development sets and extract feature vectors using TF-IDF. `(10 marks)`
* The number of features in tfidf encoded data should be 1024



<!-- grade_id: split_dataset, points: 0, grade: false, solution: true -->
```python
from sklearn.feature_extraction.text import TfidfVectorizer


def extract_features(train_df, dev_df, test_df):
    """This function takes the dataframes for training, tesing and validation and
        returns encoded features and the labels required to train the models

        X_train = [training TFIDF features]
        X_dev = [validation TFIDF features]
        X_test = [test TFIDF features]

        y_train= [lab_1, lab_2 , ..., lab_n]
        y_test= [lab_1, lab_2 , ..., lab_n]
        y_dev = [lab_1, lab_2 , ..., lab_n]
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
(train_x, train_y), (dev_x, dev_y), (test_x, test_y) = extract_features(train_df, dev_df, test_df)
```

<!-- grade_id: split_ds_funct_check, points: 5, grade: true, solution: false -->
```python
funct_text = inspect.getsource(extract_features)

funct_text_no_multiline_comment = re.sub(r'""".*?"""', '', funct_text, flags=re.DOTALL)

funct_no_comment = re.sub(r'^\s*#.*$', '', funct_text_no_multiline_comment, flags=re.MULTILINE)

n_fit = re.findall(r'\.fit.*\s', funct_no_comment)

assert len(n_fit) == 1

```

<!-- grade_id: split_ds_funct_shape, points: 5, grade: true, solution: false -->
```python
assert train_x.shape[0] == train_df.shape[0]
assert dev_x.shape[0] == dev_df.shape[0]
assert test_x.shape[0] == test_df.shape[0]


```

---
* Using these features please train a classifier using a method such Support Vector Machines (SVM) `(10 marks)`

<!-- grade_id: svm_train_a1, points: 0, grade: false, solution: true -->
```python
from sklearn.svm import SVC

def get_model(train_x, train_y):
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
model = get_model(train_x, train_y)
```

<!-- grade_id: svm_train_t1, points: 3, grade: true, solution: false -->
```python
model_type = str(type(model))
assert 'sklearn' in model_type
assert 'svm' in model_type
```

<!-- grade_id: svm_train_a2, points: 0, grade: false, solution: true -->
```python
def make_predictions_on_test(test_x, model):
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
pred_y = make_predictions_on_test(test_x, model)
```

<!-- grade_id: svm_train_t2, points: 2, grade: true, solution: false -->
```python
assert len(pred_y) == test_y.shape[0]

```

<!-- grade_id: svm_train_a3, points: 0, grade: false, solution: true -->
```python
from sklearn.metrics import precision_score, recall_score, f1_score

def calc_pre_recall_f1(pred_y, test_y):
    # YOUR CODE HERE
    raise NotImplementedError()

```

```python
prec_mac_svm, rec_mac_svm, f1_mac_svm = calc_pre_recall_f1(pred_y, test_y)
```

<!-- grade_id: svm_train_t3, points: 5, grade: true, solution: false -->
```python
assert prec_mac_svm <= 1.0
assert prec_mac_svm > 0.0

assert rec_mac_svm <= 1.0
assert rec_mac_svm > 0.0

assert f1_mac_svm <= 1.0
assert f1_mac_svm > 0.0
```

<!-- grade_id: svm_train_t4, points: 5, grade: true, solution: false -->
```python
# DO NOT DELETE !!!
```

## Task 2c

Based on the methods discussed in the lecture suggest **ONE** improvement or alternative approach that can be applied to the sentiment classification task as in Task 2a. This may be through better feature extraction, new modelling or through an alternaitve methodology.

You should implement this approach and compare the results using the evaluation from Task 2b.

Provide an **explanation** of your approach in no more than 100 words.

`(20 marks)`

## Explanation for Task

```
Put your explanation for task 2c here

```

<!-- grade_id: ml_improve_a1, points: 0, grade: false, solution: true -->
```python
def improved_predictions(train_x, train_y, test_x):
    """
    implement an improved method for task 2a:
    1. This method should take train_x, train_y used earlier and perform learning steps.
    2. These would be the implementation of the aforementioned explanation.
    3. Then it should use the test_x to return the predictions of the classification problem
    """

    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: ml_improve_t1, points: 5, grade: true, solution: false -->
```python
y_pred_improved = improved_predictions(train_x, train_y, test_x)
assert len(y_pred_improved) == test_y.shape[0]
assert all([y in set(test_y) for y in set(y_pred_improved)])
```

<!-- grade_id: ml_impove_t2, points: 5, grade: true, solution: false -->
```python
prec_mac_improv, rec_mac_improv, f1_mac_improv = calc_pre_recall_f1(y_pred_improved, test_y)
# DO NOT DELETE THIS CELL !!!
```

# 3. Span Extraction
## Task 3a


Our goal for this task is to extract the sentiment span, given the text and the given sentiment (true sentiment label) as input. The sentiment span is a subsection of the text that expresses the sentiment that classifies the text overall.

INPUT : (`text`, `sentiment`)

OUTPUT: `selected_text`

<img src="https://github.com/gauneg/gauneg.github.io/blob/main/assignment_diag.png?raw=true" alt="PLOT EXAMPLE" width="500px"/>

1. Describe in no more than 200 words a system that could be used to identify the sentiment span. You should consider the methodology potentially including how features are extracted, what models could be used and what procedures should be used to train the model. `(10 marks)`

## 200 words description in this cell (EDIT THIS)

---
2. Implement this system and apply it to the train and development splits of the dataset. `(15 marks)`

<!-- grade_id: span_extracton_impl, points: 15, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

---
## Task 3b

1. Describe an automatic metric that can be used to evaluate the task of span extraction. Implement this metric and use it to evaluate the performance of the system you developed in Task 3a. `(5 marks + 10 mark)`

### Description Goes HERE

<!-- grade_id: span_extraction_impl, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```